# *Binary Tree*

**Prerequisites**: [*Concepts of Tree and its operations*]

## Exercise 1

Write an algorithm to compute the diameter of binary tree.The diameter of the binary tree is equal to the number of nodes on 
the logest path between any two leaves in it. 
Constraints:
     
1)Check if root is None then return 0, diameter 

2)otherwise,

left_height, diameter = getDiameter(root.left, diameter)
    
right_height, diameter = getDiameter(root.right, diameter)
    
max_diameter = left_height + right_height + 1
    
diameter = max(diameter, max_diameter)
    
return max(left_height, right_height) + 1, diameter
    
Sample input:

root = Node(1)

root.left = Node(2)

root.right = Node(3)

root.left.right = Node(4)

root.right.left = Node(5)

root.right.right = Node(6)

root.right.left.left = Node(7)

root.right.left.right = Node(8)
    
Sample output:

6

#### Code Cell
Below is the code cell for **Exercise 1**. Type in your code in the cell. 

<!-- grade_id: cell-2ec62f126b24dfbd, points: 0, grade: false, solution: true -->
```python
class Node:
    def __init__(self, data, left=None, right=None):
        self.data = data
        self.left = left
        self.right = right



def getDiameter(root, diameter):

   
    ##your code here
    ###BEGIN SOLUTION
    
    if root is None:
        return 0, diameter
    left_height, diameter = getDiameter(root.left, diameter)
    right_height, diameter = getDiameter(root.right, diameter)

    max_diameter = left_height + right_height + 1


    diameter = max(diameter, max_diameter)

    return max(left_height, right_height) + 1, diameter
    
    ###END SOLUTION

def getBTDiameter(root):

    diameter = 0
    return getDiameter(root, diameter)[1]
```

#### Test Cases 
Run the below test cell to validate your code for **Exercise 1**. The cell has test cases your code should pass. Please do not modify the contents of this cell.

<!-- grade_id: cell-933483ef34527a9c, points: 0, grade: true, solution: false -->
```python
root = Node(1) 
root.left = Node(2)
root.right = Node(3)
assert getBTDiameter(root) == 3

root = Node(66)
root.left = Node(65)
root.right = Node(67)
root.left.right = Node(65.5)
root.right.left = Node(64)
assert getBTDiameter(root) == 5

print("ALL PASSED!")
```

Output:
```
ALL PASSED!
```

----
## Exercise 2


Write an algorithm to check if two binary trees are identical or not.If they have identical structure their contents are 
same.

Constraints:
    
1) Check if x is None and y is None then return True

2) Otherwisereturn (x and y) and (x.key == y.key) and \isIdentical(x.left, y.left) and isIdentical(x.right, y.right)

Sample input:

x = Node(4)

x.left = Node(2)

x.right = Node(6)

x.left.left = Node(1)

x.left.right = Node(3)

x.right.left = Node(5)

x.right.right = Node(7) 

y = Node(4)

y.left = Node(2)

y.right = Node(6)

y.left.left = Node(1)

y.left.right = Node(3)

y.right.left = Node(5)

y.right.right = Node(7) 

Sample output:

True

#### Code Cell
Below is the code cell for **Exercise 2**. Type your code in the cell.

<!-- grade_id: cell-95b9b126dd4a4a44, points: 0, grade: false, solution: true -->
```python
class Node:
    def __init__(self, key=None, left=None, right=None):
        self.key = key
        self.left = left
        self.right = right

def isIdentical(x, y):
    
    ##your code here
    ###BEGIN SOLUTION
    
    if x is None and y is None:
        return True
    return (x and y) and (x.key == y.key) and \
            isIdentical(x.left, y.left) and isIdentical(x.right, y.right)

    ###END SOLUTION
```

#### Test Cases 
Run the below test cell to validate your code for **Exercise 2**. The cell has test cases your code should pass. Please do not modify the contents of this cell.

<!-- grade_id: cell-f502e8909dcf6eb8, points: 0, grade: true, solution: false -->
```python
x=Node(10)
x.left=Node(9)
x.right=Node(11)

y=Node(10)
y.left=Node(9)
y.right=Node(11)

assert isIdentical(x, y) == True

x = Node(15)
x.left = Node(10)
x.right = Node(20)
x.left.left = Node(8)
x.left.right = Node(12)
x.right.left = Node(16)
x.right.right = Node(25)

y = Node(15)
y.left = Node(10)
y.right = Node(20)
y.left.left = Node(8)
y.left.right = Node(12)
y.right.left = Node(16)
y.right.right = Node(25)
    
assert isIdentical(x, y) == True

print("ALL PASSED!")
```

Output:
```
ALL PASSED!
```

----
## Exercise 3


Write an efficient algorithm to comput the height of binarytree.The height or depth of the tree is number of edges or nodes on 
the longest path from root node to leaf node. 

To complete the height function: 

1)Check if root is None the return 0
       
2)Otherwise return 1 + max(height(root.left), height(root.right))

Sample Input:
    
root = Node(4)
root.left = Node(2)
root.right = Node(6)
root.left.left = Node(1)
root.left.right = Node(3)
root.right.left = Node(5)
root.right.right = Node(7) 

Sample Output:
    
3

#### Code Cell
Below is the code cell for **Exercise 3**. Type your code in the cell.

<!-- grade_id: cell-70a9b90a2f9fd107, points: 0, grade: false, solution: true -->
```python
class Node:
    
    def __init__(self, key=None, left=None, right=None):
        self.key = key
        self.left = left
        self.right = right

def height(root):
    
    ##your code here
    ###BEGIN SOLUTION
    
    if root is None:
        return 0
    return 1 + max(height(root.left), height(root.right))
    
    ###END SOLUTION
    
if __name__ == '__main__':

    
    
    print("The height of the binary tree is", height(root))
```

Output:
```
The height of the binary tree is 3
```

#### Test Cases 
Run the below test cell to validate your code for **Exercise 3**. The cell has test cases your code should pass. Please do not modify the contents of this cell.

<!-- grade_id: cell-33d4a01fd430ba7b, points: 0, grade: true, solution: false -->
```python
root = Node(60)
root.left = Node(35)
root.right = Node(77)
assert height(root) == 2

root = Node(35)
root.left = Node(10)
root.right = Node(36)
root.left.left = Node(8)
root.left.right = Node(12)
assert height(root) == 3

root = Node(15)
root.left = Node(10)
root.right = Node(20)
root.left.left = Node(8)
root.left.right = Node(12)
root.right.left = Node(16)
root.right.right = Node(25)
assert height(root) == 3
    
print("ALL PASSED!")
```

Output:
```
ALL PASSED!
```

----
## Exercise 4

Write a Python program to check whether a given a binary tree is a valid binary search tree (BST) or not.

Let a binary search tree (BST) is defined as follows:
    
The left subtree of a node contains only nodes with keys less than the node's key.

The right subtree of a node contains only nodes with keys greater than the node's key.

Both the left and right subtrees must also be binary search trees.

Constraint: 

1)Perform while root or stack and while root

2)stack.append(root) and assign root.left to root and stack.pop to root.
   
3)Check if prev and root.val <= prev.val then return False assign root to prev and root.right to root then return True
       
Example 1:
  
Sample Input:

2

1

3

Sample Output:

True

Example 2:
  
Sample Input:

1

2

3

Sample Output:

False

#### Code Cell
Below is the code cell for **Exercise 4**. Type your code in the cell.

<!-- grade_id: cell-5d0ea47ca640c947, points: 0, grade: false, solution: true -->
```python
class TreeNode(object):
    
    def __init__(self, x):
        
        ##your code here
        ###BEGIN SOLUTION
        
        self.val = x
        self.left = None
        self.right = None
        
        ###END SOLUTION

def is_BST(root):
    stack = []
    prev = None
    ##your code here
    ###BEGIN SOLUTION
    
    while root or stack:
        while root:
            stack.append(root)
            root = root.left
        root = stack.pop()
        if prev and root.val <= prev.val:
            return False
        prev = root
        root = root.right
    return True
    
    ###END SOLUTION
```

#### Test Cases
Run the below test cell to validate your code for **Exercise 4**. The cell has test cases your code should pass. Please do not modify the contents of this cell.

<!-- grade_id: cell-833f0ae200c2a3ea, points: 0, grade: true, solution: false -->
```python
root = TreeNode(1)  
root.left = TreeNode(2)  
root.right = TreeNode(5) 
assert is_BST(root) == False

root = TreeNode(6)  
root.left = TreeNode(1)  
root.right = TreeNode(10) 
assert is_BST(root) == True

root = TreeNode(2)  
root.left = TreeNode(1)  
root.right = TreeNode(3) 
assert is_BST(root) == True

root = TreeNode(1)  
root.left = TreeNode(2)  
root.right = TreeNode(3) 
assert is_BST(root) == False

print("ALL PASSED!")
```

Output:
```
ALL PASSED!
```

----
## Exercise 5

Write a Python program to find the kth smallest element in a given a binary search tree.

Constraints: 
    
1)perform while root or stack and while root then stack.append(root) and assign root.left,stack.pop to root and decrement k by 1.
 
2)Check if k equals 0 then break and assign root.right to root.

3)Then return root.val.

#### Code Cell
Below is the code cell for **Exercise 5**. Type your code in the cell.

<!-- grade_id: cell-658b3178f675dcb0, points: 0, grade: false, solution: true -->
```python
class TreeNode(object):
    
    def __init__(self, x):
        self.val = x
        self.left = None
        self.right = None

def kth_smallest(root, k):
    stack = []
    
    ##your code here
    ###BEGIN SOLUTION
    
    while root or stack:
        while root:
            stack.append(root)
            root = root.left
        root = stack.pop()
        k -= 1
        if k == 0:
            break
        root = root.right
    return root.val
    
    ###END SOLUTION
```

#### Test Cases
Run the below test cell to validate your code for **Exercise 5**. The cell has test cases your code should pass. Please do not modify the contents of this cell.

<!-- grade_id: cell-6344067b20760ae0, points: 0, grade: true, solution: false -->
```python
root = TreeNode(8)  
root.left = TreeNode(5)  
root.right = TreeNode(14) 
root.left.left = TreeNode(4)  
root.left.right = TreeNode(6) 
root.left.right.left = TreeNode(8)  
root.left.right.right = TreeNode(7)  
root.right.right = TreeNode(24) 
root.right.right.left = TreeNode(22)  

assert kth_smallest(root, 2) == 5

assert kth_smallest(root, 3) == 8

assert kth_smallest(root, 1) == 4

assert kth_smallest(root, 5) == 7

assert kth_smallest(root, 4) == 6

print("ALL PASSED!")
```

Output:
```
ALL PASSED!
```