# *Circular Linked List*

**Prerequisites**: [*Concepts of linked list,Circular linked list and its operations*]

## Exercise 1

Write a function to insert a node into the circular linked list,where push is the function to insert the node into the circular linked list.

Sample input:

1

2

3

4
    
Sample output:

[1, 2, 3, 4]

#### Code Cell
Below is the code cell for **Exercise 1**. Type in your code in the cell. 

<!-- grade_id: cell-2ec62f126b24dfbd, points: 0, grade: false, solution: true -->
```python
class Node:
    
    def __init__(self,data):
        self.data=data
        self.next=None
        self.prev=None
        
class linkedlist:
    
    ##your code here
    ###BEGIN SOLUTION
    
    def __init__(self):
        self.head=None
        self.tail=None
        
    def push(self,data):
        if(self.head==None):
            nn=Node(data)
            self.head=self.tail=nn
            self.head.next=self.tail
            self.head.prev=self.tail
        else:
            nn=Node(data)
            self.tail.next = nn
            nn.prev = self.tail    
            self.tail = nn
            self.head.prev=self.tail
            self.tail.next = self.head
        return True
    
    def display(self):
        tmp=self.head
        co=0
        a=[]
        if(self.head is None):
            print("The list is empty")
        else:
            a.append(tmp.data)
            while(tmp.next!=self.head):
                tmp=tmp.next
                a.append(tmp.data)
            return(a)
        
    ###END SOLUTION
```

#### Test Cases 
Run the below test cell to validate your code for **Exercise 1**. The cell has test cases your code should pass. Please do not modify the contents of this cell.

<!-- grade_id: cell-933483ef34527a9c, points: 0, grade: true, solution: false -->
```python
ll=linkedlist()
ll.push(1)
ll.push(2)
ll.push(3)
ll.push(4)
assert ll.display() == [1, 2, 3, 4]

ll=linkedlist()
ll.push(99)
ll.push(78)
ll.push(63)
ll.push(41)
ll.push(32)
assert ll.display() == [99, 78, 63, 41, 32]

print("ALL PASSED!")
```

Output:
```
ALL PASSED!
```

----
## Exercise 2


Write a function to print the length of the linked list,Where length of the linked list is evaluated by the function length() below
 
Sample input:

1

2

3

4

Sample output:

4

#### Code Cell
Below is the code cell for **Exercise 2**. Type your code in the cell.

<!-- grade_id: cell-95b9b126dd4a4a44, points: 0, grade: false, solution: true -->
```python
def length(ll):
    
    ##your code here
    ###BEGIN SOLUTION
    
    tmp=ll.head
    co=1
    while(tmp.next!=ll.head):
        co+=1
        tmp=tmp.next
    return co

    ###END SOLUTION
```

#### Test Cases 
Run the below test cell to validate your code for **Exercise 2**. The cell has test cases your code should pass. Please do not modify the contents of this cell.

<!-- grade_id: cell-f502e8909dcf6eb8, points: 0, grade: true, solution: false -->
```python
#Before running the test cases run excercise 1 code cell,

ll=linkedlist()
ll.push(1)
ll.push(2)
ll.push(3)
ll.push(4)
assert length(ll) == 4

ll=linkedlist()
ll.push(100)
ll.push(1)
ll.push(0)
ll.push(200)
ll.push(500)
assert length(ll) == 5

print("ALL PASSED!")
```

Output:
```
ALL PASSED!
```

----
## Exercise 3


Write a function to print the position of the given node.The position of the node is evaluated by the function pos() given below.
 
Note:The position of the circular linked list starts from 1.
 
Sample input:

2

1

2

3

4

Sample output:

2

Explanation:

The position starts from 1st position,2nd position and so on.

Hence the position of val 2 is at the 2nd position and hence the output. 

#### Code Cell
Below is the code cell for **Exercise 3**. Type your code in the cell.

<!-- grade_id: cell-70a9b90a2f9fd107, points: 0, grade: false, solution: true -->
```python
def pos(ll,val):
    
    ##your code here
    ###BEGIN SOLUTION
    
    tmp=ll.head
    co=1
    if(tmp.next.data==ll.head.data):
        return co
    else:
        while(tmp.next!=ll.head):
            if(tmp.data==val):
                return co
            if(tmp.data==ll.tail.prev.data):
                return co+1
            tmp=tmp.next
            co+=1
            
    ###END SOLUTION
```

#### Test Cases 
Run the below test cell to validate your code for **Exercise 3**. The cell has test cases your code should pass. Please do not modify the contents of this cell.

<!-- grade_id: cell-33d4a01fd430ba7b, points: 0, grade: true, solution: false -->
```python
ll=linkedlist()
val=2
ll.push(1)
ll.push(2)
ll.push(3)
ll.push(4)
assert pos(ll,val) == 2

ll=linkedlist()
val=500
ll.push(50)
ll.push(100)
ll.push(200)
ll.push(500)
assert pos(ll,val) == 4

print("ALL PASSED!")
```

Output:
```
ALL PASSED!
```

----
## Exercise 4

Write a function to delete all the odd numbers in the circular linked list.The Deletion is done by the function delete() given below.After deletion only the even numbered nodes are displayed.These function contains 4 functions otherthan the delete function and they are,

 delatbeg() -> This function is used to delete the headnode.
 
 delatend() -> This function is used to delete the tailnode.
 
 delatmid() -> This function is used to delete the node otherthan in head and tail (i.e).in the middle of the list.
 
 getpos()   -> This function is used to get the position of a particular pointer.
 
Sample input:
 
1

2

3

4

Sample output:

[2, 4]

Explanation:

The list contains [1, 2, 3, 4] the odd elements in the list are 1, 3 so after deleting 1, 3 the resultant list is [2, 4].

#### Code Cell
Below is the code cell for **Exercise 4**. Type your code in the cell.

<!-- grade_id: cell-5d0ea47ca640c947, points: 0, grade: false, solution: true -->
```python
def delatend(ll):       
    if(ll.head.next==ll.head):
        ll.head=ll.tail=None
        tmp=None
    else:
        ll.tail=ll.tail.prev
        ll.tail.next=ll.head
        ll.head.prev=ll.tail
        tmp=None
        
def getpos(ll,mlt):
    tmp=ll.head
    co=1
    if(tmp.next==ll.head):
        return co
    else:
        while(tmp.next!=ll.head):
            if(tmp==mlt):
                return co
            if(tmp==ll.tail.prev):
                return co+1
            tmp=tmp.next
            co+=1
            
def delatbeg(ll):
    if(ll.head.next==ll.head):
        ll.head=ll.tail=None
        tmp=None
    else:
        ll.head=ll.head.next
        ll.head.prev=ll.tail
        ll.tail.next=ll.head
        tmp=None
        
def delatmid(ll):
    kl=ll
    lo=ll.prev
    ps=ll.next
    lo.next=ll.next
    ps.prev=lo
    kl=None
    
def delete(ll):
    
    ##your code here
    ###BEGIN SOLUTION
    tmp=ll.head
    fl=0
    if(tmp is None):
        fl=0
    else:
        while(tmp.next!=ll.head):
            ty=getpos(ll,tmp)
            if(tmp.data%2!=0):
                fl+=1
                if(ty==1):
                    delatbeg(ll)
                elif(ty==length(ll)):
                    delatend(ll)
                else:
                    delatmid(tmp)
            tmp=tmp.next
        if(tmp.data%2!=0):
            delatend(ll)
    if(fl>0):
        return True
    else:
        return False
    
    ###END SOLUTION
```

#### Test Cases
Run the below test cell to validate your code for **Exercise 4**. The cell has test cases your code should pass. Please do not modify the contents of this cell.

<!-- grade_id: cell-833f0ae200c2a3ea, points: 0, grade: true, solution: false -->
```python
ll=linkedlist()
ll.push(1)
ll.push(2)
ll.push(3)
ll.push(4)
delete(ll)
assert ll.display() == [2, 4]

ll=linkedlist()
ll.push(0)
ll.push(1)
ll.push(2)
ll.push(3)
ll.push(4)
ll.push(5)
ll.push(6)
delete(ll)
ll.display() == [0, 2, 4, 6]

print("ALL PASSED!")
```

Output:
```
ALL PASSED!
```

----
## Exercise 5

Write a function to split the given circular linked list into 2 circular linked list one containing only the even numbers and other with odd numbers.The splitting is done by the function split() given below.
 
Note:The 2 circular linked list is separated by using "end" while displaying the linked list
 
Sample input:

1

2

3

4
    
Sample output:

[1, 3]

[2, 4]

Explanation:

The list contains [1, 2, 3, 4],here we split even and odd elements into two lists i.e,[1, 3] and [2, 4] hence the output.

#### Code Cell
Below is the code cell for **Exercise 5**. Type your code in the cell.

<!-- grade_id: cell-658b3178f675dcb0, points: 0, grade: false, solution: true -->
```python
def split(ll,ev,od):
    
    ##your code here
    ###BEGIN SOLUTION
    
    tmp=ll.head
    evfl=0
    odfl=0
    while(tmp.next!=ll.head):
        if(tmp.data%2!=0):
            od.push(tmp.data)
            odfl+=1
        else:
            ev.push(tmp.data)
            evfl+=1
        tmp=tmp.next
    if(ll.tail.data%2!=0):
        od.push(ll.tail.data)
        odfl+=1
    else:
        ev.push(ll.tail.data)
        evfl+=1
    if((odfl>0) and (evfl>0)):
        return True
    else:
        return False
    
    ###END SOLUTION
```

#### Test Cases
Run the below test cell to validate your code for **Exercise 5**. The cell has test cases your code should pass. Please do not modify the contents of this cell.

<!-- grade_id: cell-6344067b20760ae0, points: 0, grade: true, solution: false -->
```python
ll=linkedlist()
ll.push(1)
ll.push(2)
ll.push(3)
ll.push(4)
od=linkedlist()
ev=linkedlist()
split(ll,ev,od)
assert ev.display() == [2, 4]

assert od.display() == [1, 3]

ll=linkedlist()
ll.push(0)
ll.push(3)
ll.push(5)
ll.push(7)
ll.push(8)
od=linkedlist()
ev=linkedlist()
split(ll,ev,od)
assert ev.display() == [0, 8]

assert od.display() == [3, 5, 7]

print("ALL PASSED!")
```

Output:
```
ALL PASSED!
```