# *Circular Queue*

**Prerequisites**: [*Concepts of Queue and its operations*]

## Exercise 1

Write a function to insert a node in the circular queue and display the result

Note:The value of the rear is incremented in enqueue operation

Sample input:

1

2

3

4

Sample output:
   
[1, 2, 3, 4]

#### Code Cell
Below is the code cell for **Exercise 1**. Type in your code in the cell. 

<!-- grade_id: cell-2ec62f126b24dfbd, points: 0, grade: false, solution: true -->
```python
class Node:
    
    def __init__(self,data):
        
        self.data=data
        self.previous=None
        self.next=None
        
class Circularqueue:
    
    ##your code here
    ###BEGIN SOLUTION
    
    def __init__(self):
        self.head=None
        self.tail=None
        self.rear=-1
        self.front=0
        
    def enqueue(self,data):
        if(self.head==None):
            nn=Node(data)
            self.head=self.tail=nn
            self.head.next=self.tail
            self.head.prev=self.tail
        else:
            newNode=Node(data)
            self.tail.next = newNode;   
            newNode.previous = self.tail;    
            self.tail = newNode;
            self.head.previous=self.tail
            self.tail.next = self.head;  
        self.rear+=1
        
    def display(self):
        a=[]
        tmp=self.head
        if(self.head is None):
            print("The list is empty")
        else:
            a.append(tmp.data)
            while(tmp.next!=self.head):
                tmp=tmp.next
                a.append(tmp.data)
            return(a)
                
    ###END SOLUTION
```

#### Test Cases 
Run the below test cell to validate your code for **Exercise 1**. The cell has test cases your code should pass. Please do not modify the contents of this cell.

<!-- grade_id: cell-933483ef34527a9c, points: 0, grade: true, solution: false -->
```python
l=Circularqueue()
l.enqueue(4)
l.enqueue(3)
l.enqueue(2)
l.enqueue(1)
assert l.rear == 3

assert l.front == 0

assert l.display() == [4, 3, 2, 1]

l=Circularqueue()
l.enqueue(109)
l.enqueue(306)
l.enqueue(201)
l.enqueue(110)
l.enqueue(409)
l.enqueue(0)
assert l.rear == 5

assert l.front == 0

assert l.display() == [109, 306, 201, 110,409, 0]

print("ALL PASSED!")
```

Output:
```
ALL PASSED!
```

----
## Exercise 2


Write a function to print the position of a node in the circular queue.

Note:function getpos gets the object of circular queue and the node to delete as the parameters.it returns a list containing the position at the 0th index and the pointer at the 1st index.

Sample input:

1

2

3

4

1
    
Sample output:

1

Explanation:

The elements are '1, 2, 3, 4' if the pos of '1' is to be it returns '1'.Because '1' is present at first place.

#### Code Cell
Below is the code cell for **Exercise 2**. Type your code in the cell.

<!-- grade_id: cell-95b9b126dd4a4a44, points: 0, grade: false, solution: true -->
```python
def getpos(ll,mlt):
    
    ##your code here
    ###BEGIN SOLUTION
    
    tmp=ll.head
    co=1
    fl=0
    if(tmp.next==ll.head):
        msk=[]
        msk.append(co,tmp)
        return msk
    else:
        while(tmp.next!=ll.head):
            psk=[]
            if(tmp.data==mlt):
                psk.append(co)
                psk.append(tmp)
                return psk
            if(tmp==ll.tail.previous):
                psk.append(co+1)
                psk.append(tmp.next)
                return psk
            tmp=tmp.next
            co+=1
            
    ###END SOLUTION
```

#### Test Cases 
Run the below test cell to validate your code for **Exercise 2**. The cell has test cases your code should pass. Please do not modify the contents of this cell.

<!-- grade_id: cell-f502e8909dcf6eb8, points: 0, grade: true, solution: false -->
```python
l=Circularqueue()
l.enqueue(1)
l.enqueue(2)
l.enqueue(3)
l.enqueue(4)
assert getpos(l,1)[0] == 1

l=Circularqueue()
l.enqueue(99)
l.enqueue(89)
l.enqueue(39)
l.enqueue(49)
l.enqueue(19)
assert getpos(l,19)[0] == 5

print("ALL PASSED!")
```

Output:
```
ALL PASSED!
```

----
## Exercise 3


Write a function to print the length of the circular queue.

Sample input:

1

2

3

4
 
Sample output:

4

#### Code Cell
Below is the code cell for **Exercise 3**. Type your code in the cell.

<!-- grade_id: cell-70a9b90a2f9fd107, points: 0, grade: false, solution: true -->
```python
def length(ll):
    
    ##your code here
    ###BEGIN SOLUTION
    
    tmp=ll.head
    co=1
    while(tmp.next!=ll.head):
        co+=1
        tmp=tmp.next
    return co

    ###END SOLUTION
```

#### Test Cases 
Run the below test cell to validate your code for **Exercise 3**. The cell has test cases your code should pass. Please do not modify the contents of this cell.

<!-- grade_id: cell-33d4a01fd430ba7b, points: 0, grade: true, solution: false -->
```python
l=Circularqueue()
l.enqueue(1)
l.enqueue(2)
l.enqueue(3)
l.enqueue(4)
assert length(l) == 4

l=Circularqueue()
l.enqueue(0)
l.enqueue(19)
l.enqueue(56)
l.enqueue(32)
l.enqueue(87)
assert length(l) == 5

l=Circularqueue()
l.enqueue(0)
assert length(l) == 1

print("ALL PASSED!")
```

Output:
```
ALL PASSED!
```

----
## Exercise 4

Write a function to delete the node by the given node and print the result.The node to be deleted is stored in "todelete" variable.where parameters in delete is l--> circular queue object,ms[0] is position of the variable to delete,ms[1] is the pointer that is pointing to the todelete variable.

delatbeg(ll) --> function to delete the head of the circular queue.

delatend(ll) --> function to delete the tail of the circular queue.

delatmid(ll) --> function to delete other than head and tail.

Note:The value of the front is incremented in dequeue operation,if the queue is empty the output is printed as "The doublylinkedqueue is empty"

Sample input:

1

1

2

3

4

5

Sample output:

[2, 3, 4, 5]

#### Code Cell
Below is the code cell for **Exercise 4**. Type your code in the cell.

<!-- grade_id: cell-5d0ea47ca640c947, points: 0, grade: false, solution: true -->
```python
def delatbeg(ll):
    if(ll.head.next==ll.head):
        ll.head=ll.tail=None
        tmp=None
    else:
        ll.head=ll.head.next
        ll.head.previous=ll.tail
        ll.tail.next=ll.head
        tmp=None
        
def delatend(ll):
    if(ll.head.next==ll.head):
        ll.head=ll.tail=None
        tmp=None
    else:
        ll.tail=ll.tail.previous
        ll.tail.next=ll.head
        ll.head.previous=ll.tail
        tmp=None
        
def delatmid(ll):
    kl=ll
    lo=ll.previous
    ps=ll.next
    lo.next=ll.next
    ps.previous=lo
    kl=None
    
def delete(ll,ty,key):
    
    ##your code here
    ###BEGIN SOLUTION

    tmp=ll.head
    if(tmp is None):
        print("The list is empty")
    else:
        if(ty==1):
            delatbeg(ll)
        elif(ty==length(ll)):
            delatend(ll)
        else:
            uy=getpos(ll,key)
            delatmid(uy[1])
        ll.front+=1
        
    ###END SOLUTION
```

#### Test Cases
Run the below test cell to validate your code for **Exercise 4**. The cell has test cases your code should pass. Please do not modify the contents of this cell.

<!-- grade_id: cell-833f0ae200c2a3ea, points: 0, grade: true, solution: false -->
```python
l=Circularqueue()
l.enqueue(1)
l.enqueue(2)
l.enqueue(3)
l.enqueue(4)
l.enqueue(5)
todelete=1
ms=getpos(l,todelete)
delete(l,ms[0],ms[1].data)
assert l.display() == [2, 3, 4, 5]

l=Circularqueue()
l.enqueue(0)
l.enqueue(2)
l.enqueue(4)
todelete=0
ms=getpos(l,todelete)
delete(l,ms[0],ms[1].data)
assert l.display() == [2, 4]

print("ALL PASSED!")
```

Output:
```
ALL PASSED!
```

----
## Exercise 5

Write a function to check whether the circular queue is empty or not.

Sample input:

"Blank"
    
Sample output:

True

Explanation:

It checks wheather there is no elements in the Queue.If their is no elements it returns 'True' else 'False'.

#### Code Cell
Below is the code cell for **Exercise 5**. Type your code in the cell.

<!-- grade_id: cell-658b3178f675dcb0, points: 0, grade: false, solution: true -->
```python
def isempty(ll):
    
    ##your code here
    ###BEGIN SOLUTION
    
    if(ll.front == ll.rear+1):
        return True
    else:
        return False
    
    ###END SOLUTION
```

#### Test Cases
Run the below test cell to validate your code for **Exercise 5**. The cell has test cases your code should pass. Please do not modify the contents of this cell.

<!-- grade_id: cell-6344067b20760ae0, points: 0, grade: true, solution: false -->
```python
lm=Circularqueue()
assert isempty(lm) == True

lm=Circularqueue()
lm.enqueue(1)
assert isempty(lm) == False

print("ALL PASSED!")
```

Output:
```
ALL PASSED!
```