# *Doubly_Linked_List*

**Prerequisites**: [*Concepts of Linked list,doubly linked list and its operations*]

## Exercise 1

Write an efficient program to perform linked list operations i.e,to check wheather the list is empty,inserting a node to the head and tail of linked list and to delete a node at the head and tail of the linked list.  

Sample input:

inserthead(1)
inserthead(2)
inserttail(3)
inserttail(4)

Sample output:

[2, 1, 3, 4]
    
Explanation:

After step1: 1

After step2: 2 1

After step3: 2 1 3

After step4: 2 1 3 4
 
The resulting list is [2, 1, 3, 4]

#### Code Cell
Below is the code cell for **Exercise 1**. Type in your code in the cell. 

<!-- grade_id: cell-2ec62f126b24dfbd, points: 0, grade: false, solution: true -->
```python
class Node:
    
    def __init__(self,data):
        self.data=data
        self.previous=None
        self.next=None
        
class linkedlist:
    
    #your code here
    ###BEGIN SOLUTION
    
    def __init__(self):
        self.head=None
        self.tail=None
        
    def isempty(self):
        if(self.head==None):
            return True
        else:
            return False
        
    def inserthead(self,key):
        if(self.isempty()):
            self.head=self.tail=Node(key)
        else:
            tmp=Node(key)
            self.head.previous=tmp
            tmp.next=self.head
            self.head=tmp
            
    def length(self):
        tmp=self.head
        co=0
        while(tmp is not None):
            co+=1
            tmp=tmp.next
        return co
    
    def inserttail(self,key):
        tmp=Node(key)
        self.tail.next=tmp
        tmp.previous=self.tail
        self.tail=tmp
        
    def deletehead(self):
        temp = self.head
        self.head = self.head.next 
        if self.head is None:
            self.tail = None 
        else:
            self.head.previous = None
        return temp
    
    def deletetail(self):
        temp = self.tail
        if self.tail is None:
            self.head = None  
        elif self.tail.previous == None:
            self.tail = None
            self.head = None
        else:
            self.tail = self.tail.previous
            self.tail.next = None 
            
    def display(self):
        tmp=self.head
        co=0
        a=[]
        while tmp is not None:
            a.append(tmp.data)
            tmp=tmp.next
            co+=1
        if(co>0):
            return(a)
        
    ###END SOLUTION
```

#### Test Cases 
Run the below test cell to validate your code for **Exercise 1**. The cell has test cases your code should pass. Please do not modify the contents of this cell.

<!-- grade_id: cell-933483ef34527a9c, points: 0, grade: true, solution: false -->
```python
ll=linkedlist()
ll.inserthead(1)
ll.inserthead(2)
ll.inserttail(3)
ll.inserttail(4)

assert ll.display() == [2, 1, 3, 4]

ll.deletehead()
ll.deletetail()
ll.deletetail()
ll.deletetail()

assert ll.isempty() == True

ll.inserthead(7)
assert ll.display() == [7]

print("ALL PASSED!")
```

Output:
```
ALL PASSED!
```

----
## Exercise 2


Write a function to check whether the nodes in the linked list are palindrome or not?

Sample Input:

R

A

D

A

R
    
Sample output:

True
    
Explanation:

if you reverse the input "RADAR" you will get the same "RADAR" so it is palindrome and returns 'True'.


#### Code Cell
Below is the code cell for **Exercise 2**. Type your code in the cell.

<!-- grade_id: cell-95b9b126dd4a4a44, points: 0, grade: false, solution: true -->
```python
class Node:
    
    def __init__(self,data):
        self.data=data
        self.previous=None
        self.next=None
        
class linkedlist:
    
    ##your code here
    ###BEGIN SOLUTION
    
    def __init__(self):
        self.head=None
        self.tail=None
        
    def push(self,data):
        if(self.head==None):
            self.head=self.tail=Node(data)
            self.tail.previous=None
            self.tail.next=None
        else:
            newNode=Node(data)
            self.tail.next = newNode;   
            newNode.previous = self.tail;    
            self.tail = newNode;
            self.tail.next = None;  
            
    def length(self):
        tmp=self.head
        co=0
        while(tmp is not None):
            co+=1
            tmp=tmp.next
        return co
    
    def ispalindrome(self,k):
        co=0
        var1=self.head
        var2=self.tail
        while(co<k):
            if(var1.data==var2.data):
                var1=var1.next
                var2=var2.previous
                co+=1
            else:
                return False
                break
        if(co==k):
            return True
        
    ###END SOLUTION
```

#### Test Cases 
Run the below test cell to validate your code for **Exercise 2**. The cell has test cases your code should pass. Please do not modify the contents of this cell.

<!-- grade_id: cell-f502e8909dcf6eb8, points: 0, grade: true, solution: false -->
```python
l=linkedlist()
l.push("R")
l.push("A")
l.push("D")
l.push("A")
l.push("R")
k=l.length()//2
assert l.ispalindrome(k) == True

l=linkedlist()
l.push("M")
l.push("A")
l.push("D")
l.push("A")
l.push("M")
k=l.length()//2
assert l.ispalindrome(k) == True

l=linkedlist()
l.push("A")
l.push("D")
l.push("A")
l.push("M")
k=l.length()//2
assert l.ispalindrome(k) == False

print("ALL PASSED!")
```

Output:
```
ALL PASSED!
```

----
## Exercise 3


Write a function to rotate the linked list.

Sample input:

1

4

3

2

1

9

11
    
Sample Output:
 
[3, 2, 1, 9, 11, 4]

Explanation:

when the inp = 1, every element in the list rotates once.

#### Code Cell
Below is the code cell for **Exercise 3**. Type your code in the cell.

<!-- grade_id: cell-70a9b90a2f9fd107, points: 0, grade: false, solution: true -->
```python
class Node:
    
    def __init__(self,data):
        self.data=data
        self.previous=None
        self.next=None
        
class linkedlist:
    
    ##your code here
    ###BEGIN SOLUTION
    
    def __init__(self):
        self.head=None
        self.tail=None
        
    def push(self,data):
        if(self.head==None):
            self.head=self.tail=Node(data)
            self.tail.previous=None
            self.tail.next=None
        else:
            newNode=Node(data)
            self.tail.next = newNode;   
            newNode.previous = self.tail;    
            self.tail = newNode;
            self.tail.next = None; 
            
    def length(self):
        tmp=self.head
        co=0
        while(tmp is not None):
            co+=1
            tmp=tmp.next
        return co
    
    def rotate(self,n):
        current = self.head;
        if(n>=self.length()):
            return False
        for i in range(0,n-1):
            current=current.next
        self.tail.next=self.head
        self.head=current.next
        current.previous=None
        self.tail=current
        self.tail.next=None
        return True
    
    def display(self):
        tmp=self.head
        a=[]
        while tmp is not None:
            a.append(tmp.data)
            tmp=tmp.next
        return(a)
    
    ###END SOLUTION
```

#### Test Cases 
Run the below test cell to validate your code for **Exercise 3**. The cell has test cases your code should pass. Please do not modify the contents of this cell.

<!-- grade_id: cell-33d4a01fd430ba7b, points: 0, grade: true, solution: false -->
```python
l=linkedlist()
l.push(4)
l.push(3)
l.push(2)
l.push(1)
l.push(9)
l.push(11)
inp=1
assert l.rotate(inp) == True

assert l.display() == [3, 2, 1, 9, 11, 4]

l=linkedlist()
l.push(1)
l.push(2)
l.push(3)
inp=5
assert l.rotate(inp) == False

print("ALL PASSED!")
```

Output:
```
ALL PASSED!
```

----
## Exercise 4

Write a function to check whether the given data in linked list is sorted or not.

Sample input:

1

2

3

4

Sample output:

True

Explanation:

All the elements in the list are in ascending order and hence it is sorted.Therefore returns 'True'.

#### Code Cell
Below is the code cell for **Exercise 4**. Type your code in the cell.

<!-- grade_id: cell-5d0ea47ca640c947, points: 0, grade: false, solution: true -->
```python
class Node:
    
    def __init__(self,data):
        self.data=data
        self.previous=None
        self.next=None
        
class linkedlist:
    
    def __init__(self):
        self.head=None
        self.tail=None
    
    ##your code here
    ###BEGIN SOLUTION
    
    def push(self,data):
        if(self.head==None):
            self.head=self.tail=Node(data)
            self.tail.previous=None
            self.tail.next=None
        else:
            newNode=Node(data)
            self.tail.next = newNode;   
            newNode.previous = self.tail;    
            self.tail = newNode;
            self.tail.next = None;
            
    def length(self):
        tmp=self.head
        co=0
        while(tmp is not None):
            co+=1
            tmp=tmp.next
        return co
    
    def checksorted(self):
        tmp=self.head
        co=0
        while(co<self.length()-1):
            if(tmp.data<=tmp.next.data):
                tmp=tmp.next
                co+=1
            else:
                return False
                break
        if(co==self.length()-1):
            return True
        
    ###END SOLUTION
```

#### Test Cases
Run the below test cell to validate your code for **Exercise 4**. The cell has test cases your code should pass. Please do not modify the contents of this cell.

<!-- grade_id: cell-833f0ae200c2a3ea, points: 0, grade: true, solution: false -->
```python
l=linkedlist()
l.push(1)
l.push(2)
l.push(3)
l.push(4)
assert l.checksorted() == True

l=linkedlist()
l.push(9)
l.push(2)
l.push(10)
l.push(22)
assert l.checksorted() == False

print("ALL PASSED!")
```

Output:
```
ALL PASSED!
```

----
## Exercise 5

Write a function to sort the given elements in the linked list using bubble sort algorithm.

Sample input:

5

4

3

2

Sample output: 

[2, 3, 4, 5]

#### Code Cell
Below is the code cell for **Exercise 5**. Type your code in the cell.

<!-- grade_id: cell-658b3178f675dcb0, points: 0, grade: false, solution: true -->
```python
class Node:
    
    def __init__(self,data):
        self.data=data
        self.previous=None
        self.next=None
        
class linkedlist:
    
    ##your code here
    ###BEGIN SOLUTION
    
    def __init__(self):
        self.head=None
        self.tail=None
        
    def push(self,data):
        if(self.head==None):
            self.head=self.tail=Node(data)
            self.tail.previous=None
            self.tail.next=None
        else:
            newNode=Node(data)
            self.tail.next = newNode;   
            newNode.previous = self.tail;    
            self.tail = newNode;
            self.tail.next = None;
            
    def length(self):
        tmp=self.head
        co=0
        while(tmp is not None):
            co+=1
            tmp=tmp.next
        return co
    
    def sortll(self):
        for i in range(0,self.length()):
            swap=0
            temp=self.head
            for j in range(0,self.length()-i-1):
                if(temp.data>temp.next.data):
                    sk=temp.data
                    temp.data=temp.next.data
                    temp.next.data=sk
                    temp=temp.next
                    swap=1
                else:
                    temp=temp.next
            if(swap==0):
                break
        return True
    
    def display(self):
        a=[]
        tmp=self.head
        while tmp is not None:
            a.append(tmp.data)
            tmp=tmp.next
        return(a)
    ###END SOLUTION
```

#### Test Cases
Run the below test cell to validate your code for **Exercise 5**. The cell has test cases your code should pass. Please do not modify the contents of this cell.

<!-- grade_id: cell-6344067b20760ae0, points: 0, grade: true, solution: false -->
```python
l=linkedlist()
l.push(5)
l.push(4)
l.push(3)
l.push(2)
l.sortll()
assert l.length() == 4

assert l.display() == [2, 3, 4, 5]

l=linkedlist()
l.push(19)
l.push(15)
l.push(34)
l.push(21)
l.push(100)
l.sortll()
assert l.length() == 5

assert l.display() == [15, 19, 21, 34, 100]

print("ALL PASSED!")
```

Output:
```
ALL PASSED!
```