# *Queue*

**Prerequisites**: [*Concepts of Queue and its operations, Basics of loops and Functions*]

## Exercise 1

Write a class Queue that defines the methods enqueue, dequeue and is_empty inside the class Queue.

1. The method enqueue appends data to items.

2. The method dequeue dequeues the first element in items.

3. The method is_empty returns True only if items is empty.

4. The method peek() returns the first element of the Queue.

5. The method size() returns the size of the Queue.

#### Code Cell
Below is the code cell for **Exercise 1**. Type in your code in the cell. 

<!-- grade_id: cell-2ec62f126b24dfbd, points: 0, grade: false, solution: true -->
```python
class Queue:
    
    def __init__(self):
        self.items = []
        
        ##your code here.
        ### BEGIN SOLUTION
 
    def is_empty(self):
        return self.items == []
    
    def enqueue(self, data):
        self.items.append(data)
 
    def dequeue(self):
        return self.items.pop(0)
    
    def __str__(self):
        return str(self.items)
    
    def size(self):
        return len(self.items)
    
    def peek(self):
        if self.items:
            return self.items[-1]
        
        ### END SOLUTION
```

#### Test Cases 
Run the below test cell to validate your code for **Exercise 1**. The cell has test cases your code should pass. Please do not modify the contents of this cell.

<!-- grade_id: cell-933483ef34527a9c, points: 0, grade: true, solution: false -->
```python
q = Queue()
q.enqueue(100)
assert q.is_empty() == False

assert q.peek() == 100

assert q.size() == 1

q = Queue()
q.enqueue(1)
q.enqueue(2)
q.enqueue(3)
q.enqueue(4)
q.enqueue(5)
assert q.is_empty() == False

assert q.size() == 5

print("ALL PASSED!")
```

Output:
```
ALL PASSED!
```

----
## Exercise 2


Complete the following class Queue_Two_Stacks() to implement a Queue data structure using TWO stacks thus converting LIFO to FIFO operation.

#### Code Cell
Below is the code cell for **Exercise 2**. Type your code in the cell.

<!-- grade_id: cell-95b9b126dd4a4a44, points: 0, grade: false, solution: true -->
```python
class Queue_Two_Stacks():

    def __init__(self):
        self.stack_1 = []
        self.stack_2 = []
        
    ##your code here
    ###BEGIN SOLUTION
    
    def __str__(self):
        return str(self.stack_1)

    def enqueue(self, item):
        self.stack_1.append(item)

    def dequeue(self):
        if len(self.stack_2) == 0:
            if len(self.stack_1) == 0:
                raise IndexError("Can't dequeue from empty queue!")
            while len(self.stack_1) > 0 :
                last_stack_1_item = self.stack_1.pop()
                self.stack_2.append(last_stack_1_item)
        return self.stack_2.pop()
    
    def is_empty(self):
        return self.stack_2 == []
   
    def size(self):
        return len(self.stack_1)
    
    ###END SOLUTION
```

#### Test Cases 
Run the below test cell to validate your code for **Exercise 2**. The cell has test cases your code should pass. Please do not modify the contents of this cell.

<!-- grade_id: cell-f502e8909dcf6eb8, points: 0, grade: true, solution: false -->
```python
q1=Queue_Two_Stacks()
q1.enqueue(10)
q1.enqueue(20)
q1.enqueue(30)
assert q1.dequeue() == 10

assert q1.dequeue() == 20

assert q1.dequeue() == 30

assert q1.is_empty()==True

print("ALL PASSED!")
```

Output:
```
ALL PASSED!
```

----
## Exercise 3


Complete the following list_to_queue class to implement a Queue data structure using a List.

Use put() as enqueue() method.

Use get() as dequeue() method.

The rotate() method is used to rotate the queue structure by a count specified by the user!

Example:

q2 = [1,2,3,4,5]

get() -> returns [1]

put(6) changes the queue to q2 = [1,2,3,4,5,6] 

rotate(2) changes the queue to q2 = [3,4,5,6,1,2]

#### Code Cell
Below is the code cell for **Exercise 3**. Type your code in the cell.

<!-- grade_id: cell-70a9b90a2f9fd107, points: 0, grade: false, solution: true -->
```python
class list_to_queue:
    
    def __init__(self):
        self.entries = []
        self.length = 0
        self.front = 0

    ##your code here
    ###BEGIN SOLUTION
        
    def __str__(self):
        printed = "[" + str(self.entries)[1:-1] + "]"
        return printed

    def put(self, item):
        self.entries.append(item)
        self.length = self.length + 1

    def get(self):
        self.length = self.length - 1
        dequeued = self.entries[self.front]
        self.entries = self.entries[1:]
        return dequeued
        

    def rotate(self, rotation): 
        for i in range(rotation):
            self.put(self.get())

    def get_front(self): 
        return self.entries[0]
    
    def size(self):
        return self.length
    
    ###END SOLUTION
```

#### Test Cases 
Run the below test cell to validate your code for **Exercise 3**. The cell has test cases your code should pass. Please do not modify the contents of this cell.

<!-- grade_id: cell-33d4a01fd430ba7b, points: 0, grade: true, solution: false -->
```python
q2=list_to_queue()

assert q2.get() == 0

q2.put(100)

assert q2.size() == 1

assert q2.get_front() == 1

q2.rotate(2)

print("ALL PASSED!")
```

----
## Exercise 4

Complete the class LinkedQueue to implement a queue data structure using a Linked List.

Class Node will initialize node to be added to the Queue structure.

Use put() as enqueue() method.

Use get() as dequeue() method.

#### Code Cell
Below is the code cell for **Exercise 4**. Type your code in the cell.

<!-- grade_id: cell-5d0ea47ca640c947, points: 0, grade: false, solution: true -->
```python
from typing import Any, Optional

class Node:
    def __init__(self, data: Any, next: Optional["Node"] = None):
        self.data: Any = data
        self.next: Optional["Node"] = next
            
class LinkedQueue:
   
    ##your code here
    ###BEGIN SOLUTION
    
    def __init__(self) -> None:
        self.front: Optional[Node] = None
        self.rear: Optional[Node] = None

    def is_empty(self) -> bool:
        return self.front is None
    
    def put(self, item: Any) -> None:
        node: Node = Node(item)
        if self.is_empty():
            self.front = node
            self.rear = node
        else:
            assert isinstance(self.rear, Node)
            self.rear.next = node
            self.rear = node
        
    def get(self) -> Any:
        if self.is_empty():
            raise IndexError("get from empty queue")
        else:
            assert isinstance(self.front, Node)
            node: Node = self.front
            self.front = node.next
            if self.front is None:
                self.rear = None
            return node.data
        
        ###END SOLUTION
```

#### Test Cases
Run the below test cell to validate your code for **Exercise 4**. The cell has test cases your code should pass. Please do not modify the contents of this cell.

<!-- grade_id: cell-833f0ae200c2a3ea, points: 0, grade: true, solution: false -->
```python
queue = LinkedQueue()
assert queue.is_empty()==True

queue.put(5)
queue.put(9)
queue.put('python')
assert queue.is_empty()==False

assert queue.get()==5

queue.put('algorithms')
assert queue.get()==9

assert queue.get()=='python'

assert queue.get()=='algorithms'

assert queue.is_empty()==True

print("ALL PASSED!")
```

Output:
```
ALL PASSED!
```

----
## Exercise 5

Complete the class pseudo_stack_Queue represented by a pseudo stack (represented by a list with pop and append).

#### Code Cell
Below is the code cell for **Exercise 5**. Type your code in the cell.

<!-- grade_id: cell-658b3178f675dcb0, points: 0, grade: false, solution: true -->
```python
class Pseudo_stack_Queue:
    
    def __init__(self):
        self.stack = []
        self.length = 0
        
    ##your code here
    ###BEGIN SOLUTION 

    def __str__(self):
        printed = "[" + str(self.stack)[1:-1] + "]"
        return printed

    def put(self, item):
        self.stack.append(item)
        self.length = self.length + 1

    def get(self):
        self.rotate(1)
        dequeued = self.stack[self.length - 1]
        self.stack = self.stack[:-1]
        self.rotate(self.length - 1)
        self.length = self.length - 1
        return dequeued

    def front(self):
        front = self.get()
        self.put(front)
        self.rotate(self.length - 1)
        return front
        
    def rotate(self, rotation):
        for i in range(rotation):
            temp = self.stack[0]
            self.stack = self.stack[1:]
            self.put(temp)
            self.length = self.length - 1

    def size(self):
        return self.length
    
    ###END SOLUTION
```

#### Test Cases
Run the below test cell to validate your code for **Exercise 5**. The cell has test cases your code should pass. Please do not modify the contents of this cell.

<!-- grade_id: cell-6344067b20760ae0, points: 0, grade: true, solution: false -->
```python
qu=Pseudo_stack_Queue()
assert qu.size()==0

qu.put(10)
qu.put(20)
assert qu.front()==10

assert qu.size()==2

qu.put('Hello')
assert str(qu)=="[10, 20, 'Hello']"

assert qu.get()==10

assert qu.get()==20

assert qu.get()=='Hello'

assert qu.size()==0

print("ALL PASSED!")
```

Output:
```
ALL PASSED!
```

----
## Exercise 6

Complete the following hotPotato() function to implement the Hot Potato Simulation in which,
We will implement a general simulation of Hot Potato. Our program will input a list of names and a constant, call it “num,” to be used for counting. It will return the name of the last person remaining after repetitive counting by num. After num dequeue/enqueue operations, the child at the front will be removed permanently and another cycle will begin.

Sample Input:
    
Queue: "Bill","David","Susan","Jane","Kent","Brad"
    
Rotation Count: 3

Sample Output:
    
Kent

<!-- grade_id: cell-cbda52646632cdae, points: 0, grade: false, solution: true -->
```python
def hotPotato(namelist, num):
    simqueue = Queue() ## for this function run the class Queue in Exercise 1 ##
    
    for name in namelist:
        simqueue.enqueue(name)
        
    while simqueue.size() > 1:
        
    ##your code here
    ### BEGIN SOLUTION
    
        for i in range(num):
            simqueue.enqueue(simqueue.dequeue())
        simqueue.dequeue()
    return simqueue.dequeue()

    ### END SOLUTION
```

#### Test Cases
Run the below test cell to validate your code for **Exercise 6**. The cell has test cases your code should pass. Please do not modify the contents of this cell.

<!-- grade_id: cell-42905c50023d6ea5, points: 0, grade: true, solution: false -->
```python
assert hotPotato(["Bill","David","Susan","Jane","Kent","Brad"],1) == 'Kent'

assert hotPotato(["Bill","David","Susan","Jane","Kent","Brad"],5) == 'Jane'

print("ALL PASSED!")
```

Output:
```
ALL PASSED!
```